---
title       : Day 4
description : Creating graphics

--- type:NormalExercise xp:100 skills:1 key:fb7fd9c2c9
## Intro
Before starting this hands-on practice make sure to view this lecture [Intro to R visualization] (https://www.slideshare.net/obscrivn/r-visualization-ggplot2-googlevis-plotly-igraph-overview/)


*** =instructions


*** =hint


*** =pre_exercise_code
```{r}
# Pre-load dplyr to access mtcars
library(dplyr)
# Pre-load dplyr to access plot functions
library(ggplot2)

# load mtcars:
mtcars <- data(mtcars)
```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Awesome!")
```