---
title       : Chapter 1
description : Data Frames

--- type:NormalExercise xp:50 skills:1 key:3c1cd8a7c3
## Introduction

*** =instructions


*** =hint


*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
#test_output_contains("8", incorrect_msg = "Make sure to add `3 + 5` on a new line. Do not start the line with a `#`, otherwise your R code is not executed!")
#success_msg("Awesome! See how the console shows the result of the R code you submitted? Now that you're familiar with the interface, let's get down to R business!")
```

--- type:MultipleChoiceExercise xp:50 skills:1 key:14288e1b24

*** =instructions
#- logical
#- character
#- numeric


*** =sct
```{r}
#msg1 = "Try again! Think about quotes"
#msg2 = "Well done. Proceed to the next exercise"
#msg3 = "Try again! Think about quotes"
#test_mc(correct = 2, feedback_msgs = c(msg1,msg2,msg3))
```
